// Welcome to Starklings: Cairo v1
// This exercise doesn't do anything yet but it still compiles! Cairo file getting run
// needs to have a `main` function. So this file is a valid Cairo file.
// Other exercises will require you to write Cairo code to make the exercise file compile.
// Remove "I AM NOT DONE" comment when you are ready to move to the next exercise.

fn main() {}
