// Felts supports basic math operations.
// Only accepted operators (const excluded) are: +, -, * and /

// I AM NOT DONE

// TODO
// Return the solution of (x² + x - 2) / (x - 2)


fn poly(x : felt) ->  felt {
    // FILL ME
    res   // Do not change
}

// Do not change the test function
#[test]
fn test_poly(){
    let res = poly(0);
    assert(res==0, 'Error message');
    let res = poly(0);
    assert(res==10, 'Error message');
}
