// structs1.cairo
// Address all the TODOs to make the tests pass!
// Execute `starklings hint structs1` or use the `hint` watch subcommand for a hint.

// I AM NOT DONE
#[derive(Copy, Drop)]
struct ColorStruct {
    // TODO: Something goes here
}


#[test]
fn classic_c_structs() {
    // TODO: Instantiate a classic color struct!
    // let green =

    assert(green.red == 0, 0);
    assert(green.green == 255, 0);
    assert(green.blue == 0, 0);
}
